## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def area(x0,y0,x1,y1,x2,y2):
    A = abs((x0*(y1-y2)+x1*(y2-y0)+x2*(y0-y1))/2)
    return A

def check(co):
    if len(co)!=2:
        print('bad input')
        print(len(co))
        return True


co1 = input('Enter coordinates of first vertex, e.g "x,y": ')
co1 = co1.split(",")
check(co1)
    
co2 = input('Enter coordinates of second vertex, e.g "x,y": ')
co2 = co2.split(",")
check(co2)

co3 = input('Enter coordinates of third vertex, e.g "x,y": ')
co3 = co3.split(",")
check(co3)

area(int(co1[0]),int(co1[1]),int(co2[0]),int(co2[1]),int(co3[0]),int(co3[1]))
```

Output:
```
Enter coordinates of first vertex "x,y": 0,0
Enter coordinates of second vertex (x,y): 0,1
Enter coordinates of third vertex (x,y): 1,0
0.5
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def area(my_list):
    area_add = 0
    for i in range(len(my_list)):
        if i == int(len(my_list)) - 1:
            area_add+= my_list[i][0]*my_list[0][1]-my_list[0][0]*my_list[i][1]
        
        else:
            area_add+= my_list[i][0]*my_list[i+1][1]-my_list[i+1][0]*my_list[i][1]
    return abs(area_add/2)

polygon = [(0,0),(1,0),(0,1)]
area(polygon)
```

Output:
```
0.5
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    # YOUR CODE HERE
    if len(x)!=len(y):
        print('vectors not same amount of dimensions')
        return
    
    my_list =[]
    for i in range(len(x)):
        my_list.append(x[i]+y[i])
    return my_list

for i in range(1):
    print(type(i))

```

Output:
```
<class 'int'>
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
my_dict = {"Christ's" : "CHR", "Churchill":"CHU", "Clare":"CL","St John's":"JN","Clare Hall":"CLH" }
my_list = []
reverse_dict = dict()
for key in my_dict:
    print(f"The abbreviation of {key} is {my_dict[key]}")
    my_list.append(my_dict[key])
    reverse_dict[my_dict[key]] = key
    
print(reverse_dict)
print(sorted(my_list))
```

Output:
```
The abbreviation of Christ's is CHR
The abbreviation of Churchill is CHU
The abbreviation of Clare is CL
The abbreviation of St John's is JN
The abbreviation of Clare Hall is CLH
{'CHR': "Christ's", 'CHU': 'Churchill', 'CL': 'Clare', 'JN': "St John's", 'CLH': 'Clare Hall'}
['CHR', 'CHU', 'CL', 'CLH', 'JN']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```